Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart Kernel) and then **run all cells** (in the menubar, select Run$\rightarrow$Run All Cells). Alternatively, you can use the **validate** button in the assignment list panel.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE". When you insert your Code you can remove the line `raise NotImplementedError()`. Also put your name, matriculationnumber, and collaborators below:

```python
NAME = "Li xinghan"
MATRICULATIONNUMBER = "5130032"
COLLABORATORS = ""
```

---

<img src="images/logo_ifn.svg" alt="Drawing" style="width: 256px;" align="right"/>

# Exercise 1.1: Python Basics

Welcome to the first unit of the Computer Lab Pattern Recognition where you will learn about the programming language Python. Python is a unviversal, interpreted programming language, which is essential for most machine learning and deep learning applications. The functionality can be extended by importing libraries such as PyTorch and Tensorflow which you might have heard of before. The following exercises provide some short tasks to introduce and consolidate your Python knowledge. We will start with the Python basics, followed by an introduction to NumPy, SciPy, Pandas and Matplotlib.

We use Jupyter notebooks to step through the required fundamentals of the programming language and the used libraries. Jupyter provides browser-based, interactive notebooks for programming.
Below this text you can see what is called a cell. Each notebook consists of multiple cells, which can be used to show formattted text, images and formulas or to write and programm code. You can use different kernels to interprete your code, but we stick with a Python 3 kernel. All modifications to the notebook are automatically saved every few minutes.
The first cell below this text is a code cell starting with `[ ]`. In these cells you can type your python code and execute it by simply hitting ``Shift`` + ``Enter``. You then will be placed in a new cell. Try it by typing and executing the following code:
```python
27+2
```

```python
27 + 2
```

**Warning:**
Pay attention to not open the same notebook document on many tabs: edits on different tabs can overwrite each other! To be safe, make sure you open each notebook document in only one tab. If you accidentally open a notebook twice in two different tabs, just close one of the tabs.

In the tasks of this Computer Lab, we will provide exercises using a JupyterHub together with nbgrader. The JupyterHub is installed on a remote server, which provides you with compute resources for executing notebooks. Essentially when you execute a code cell, the code will be run on the remote server. The nbgrader extension for the JupyterHub allows for interaction of lecturers with students. We from the Institute for Communications Technology have prepared a lot of exercises for you, which we make available via nbgrader as shown in the kick-off meeting and the corresponding slides. You can fetch, validate and submit your exercises in the `Nbgrader` --> `Assignment List`. Make sure to validate your assignemnts before submitting and to only submit once, when you are confident that you have solved all tasks to your satisfaction. 

### Task 1.1-A: Simple Operations (3P)

First, we will learn how to work with variables. We can create variables using statements such as `a=13` and `b=4` as shown below. Then, we can for example add them to a thrid variable `c = a + b`. You can print the value of the variable using the `print()` function. Try executing the following cell.

```python
a = 13
b = 4
c = a + b
print(c)
```

Output:
```
17
```

Other operators such as `-`, `*`, `/`, `//`, and `%` are used for subtraction, multiplication, division, integer division, and the modulo operation, respectively. Please compute all these operation (a before the operator, b after the operator) and store them in variables `d`, `e`, `f`, `g`, `h`

<!-- grade_id: 11-a-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
d = a - b
e = a * b
f = a / b
g = a // b
h = a % b
```

We will test your solutions with tests as shown below. The `assert` statement means that an error os thrown if the condition is not fulfilled. We will use these errors to automatically grade your tasks. Also, you can use the tests shown below to validate your solutions. Accordingly, please do not try to delete or modify the test cells in any way as this might impact the evaluation of your tasks. We have usually set them to an unmodifyable state so that in theory you should anyway not be able to modify them. 

**Warning:**
Please note, that you will only see a subset of the tests and that many other hidden tests will be potentially carried out in subsequent tasks. It is therefore, not a viable solution to just copy your friends' solutions. However, you are encouraged to work together and discuss your individual solutions.

<!-- grade_id: 11-a-test, points: 3, grade: true, solution: false -->
```python
assert d == 9
assert e == 52
assert f == 3.25
assert g == 3
assert h == 1
```

### Task 1.1-B: Data Types (3P)

In Python you do not define the data type when initializing a variable (as e.g., in C++) but Python rather determines the type automatically. Still, you should be aware of which possible data types there are. You can always check the data type using the `type()` function.

Please create variables `i_value`, `f_value`, `s_value`, and `b_value`, which contain arbitrary values of type Integer (int), Floating point (float), String (str), and Boolean (bool).  



<!-- grade_id: 11-b-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
i_value = 2
f_value = 2.1
s_value = "Hallo"
b_value = True
```

Please note that in the following test we have only made the two tests for the Integer and Floating point types visible. The test for String and Boolean values is not visible but will be carried out when evaluating the notebook. The following tasks follow the spirit that we will be providing enough tests for you to validate that your output/solution has the correct format/shape, but we will not provide tests which are used to prevent cheating or to check the final quality of the solution

<!-- grade_id: 11-b-test, points: 3, grade: true, solution: false -->
```python
assert type(i_value) == int
assert type(f_value) == float

```

### Task 1.1-C: Data Structures (4P)
Often times in programming not only single values are needed but collections of values. Python offers various possibilities to create those structures, e.g., lists, dictionaries, or tuples, which contain different properties. Please check the internet to find out about their exact properties and how to work with them. Afterwards, solve the following tasks:
- Reverse the order of elements in `num_list2` and store it in a variable `reverse_list2`. Hint: be careful not to modify the original list `num_list2`.
- Merge the `num_list` and `num_list2` into a list `merge_list` and keep the order of elements such that a list starting from 1 and ending at 10 is created.
- Append the number 0 to `num_list` and store the resulting list in `new_list`
- merge the two dictionaries `first_dict` and `second_dict` into a single dictionary `merge_dict` and remove the key 'bike'


<!-- grade_id: 11-c-answer, points: 0, grade: false, solution: true -->
```python
num_list = [1, 2, 3, 4, 5]
num_list2 = [6, 7, 8, 9, 10]
char_list2 = ['a', 'b', 'c']
first_dict = {'car': 12, 'boat': 26, 'plane': 8}
second_dict = {'bike': 'nice', 'jet': 'fast', 'scooter': 'flexible'}

# YOUR CODE HERE
reverse_list2 = num_list[::-1]
merge_list = num_list + num_list2
new_list = num_list.copy()
new_list.append(0)
merge_dict = first_dict.copy()
merge_dict.update(second_dict)
del merge_dict['bike']
```

<!-- grade_id: 11-c-test, points: 4, grade: true, solution: false -->
```python
assert type(reverse_list2) == list
assert type(merge_list) == list
assert type(new_list) == list
assert type(merge_dict) == dict

```

### Task 1.1-D: Functions (5P)

Code can have thousands and even millions of code lines. Accordingly, it is not feasible to put all processing blocks into a large file, but there has to be an ordering. One of the simplest way to do this are functions. They take arguments and have return values. Write the following functions:
- Write a function `return_hello_world` that returns a String `Hello World`.
- Write a function `add_values` that receives two values and adds them.
- Write a function `sum_list` that sums up all values in a list.

<!-- grade_id: 11-d-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def return_hello_world():
    return "Hello World"
def add_values(value1,value2):
    return value1 + value2
def sum_list(list1):
    sum_value=sum(list1)
    return sum_value
```

<!-- grade_id: 11-d-test, points: 5, grade: true, solution: false -->
```python
val1 = return_hello_world()
val2 = add_values(4,9)
val3 = sum_list([4,2,3])

```

### Task 1.1-E: If, Elif, and Else Conditions (5P)

Often we would like to carry out a certain statement only if a condition is met. We can do this using the `if`, `elif`, and `else` conditions. Check the internet for more information about these statements and implement the following functions:
- Implement a step function as `step` which takes a number and returns 0 if the number is smaller than 0 or 1 if the number is larger or equal to 0.
- Implement ReLU function `relu` which takes a number and returns the input value for values larger or equal to 0 and 0 for values smaller than 0.
- Implement a function `two_cond` which takes a String and return True if the string contains an 'a' as well as an 'e' and 'False' otherwise.
- Implement a function `multi_step` which takes a number and returns 0 if the number is smaller than -1, 0.5 if the number is larger or equal to -1 as well as smaller than 1, and 1 if the number is greater or equal to 1.

<!-- grade_id: 11-e-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def step(number):
    if number < 0:
        return 0
    else:
        return 1
def relu(number):
    if number>= 0:
        return number
    else:
        return 0
def two_cond(string1):
    if 'a' in string1 and 'e' in string1:
        return True
    else:
        return False
def multi_step(number):
    if number < -1:
        return 0
    elif number >= -1 and number <1:
        return 0.5
    else:
        return 1
```

<!-- grade_id: 11-e-test, points: 5, grade: true, solution: false -->
```python
out = step(2)
out2 = relu(2)
out3 = two_cond('equal')
out4 = multi_step(2)

```

### Task 1.1-F: For and While Loops (5P)

Often times we also need to loop through data structures or numbers. We can do this with `for` and `while` loops. The `for` loop goes through a fixed length data structure, while the `while` loop goes on infinitely until a condition is met. With this knowledge implement the following functions:
- A function `count_two` that counts how often the number 2 occurs in a list
- A function `match` that simultaneously loops through two lists of numbers and counts how often the values (at the same position) match. You can assume that the lists have the same length. Hint: use the `zip()` function.
- A function that `limit_add` that takes a threshold as input and determines until which number the natural numbers starting from 1 have to be added until a given threshold is reached or surpassed. E.g., if the threshold is 12, the numbers 1,2,3,4, and 5 have to be added, so the function would return 5.

<!-- grade_id: 11-f-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def count_two(list1):
    count = 0
    for num in list1:
        if num == 2:
            count+=1
    return count
def match (list1,list2):
    index = 0
    for num1,num2 in zip(list1,list2):
        if num1== num2:
            index += 1
    return index
def limit_add(threshold):
    sum = 0
    index = 1
    while sum < threshold:
        sum += index
        index += 1
    return index -1
```

<!-- grade_id: 11-f-test, points: 5, grade: true, solution: false -->
```python
out1 = count_two([2,4,3,2,5,2,1])
out2 = match([2,3,1], [2,1,3])
out3 = limit_add(12)

```

### Task 1.1-G: Import Statements (5P)

While it is nice to write python code yourself, for complex projects it is often necessary to rely on functions other people have written. For this purpose Python offers the possibility to import libraries or functions from libraries. Famous libraries are for example numpy, scipy, pandas, or, pytorch and tensorflow. Import functions that solve the following tasks (or write them yourself if you want to):
- A function `time` that returns the current time as a floating point number expressed in seconds since the epoch
- A function `inv` that returns the inverse of an arbitrary quadratic matrix 
- A function `curr_dir` that returns the current working directory of this script as a String in the format /absolute/path/to/current/directory

<!-- grade_id: 11-g-answer, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import time as t
import numpy as np
import os

def time():
    time1=t.time()
    return time1

def inv(matrix):
    inverse = np.linalg.inv(matrix)
    return inverse

def curr_dir():
    curr_directory = os.getcwd()
    return curr_directory
    
    
```

<!-- grade_id: 11-g-test, points: 5, grade: true, solution: false -->
```python
out1 = time()
out2 = inv([[1,5],[2,1]])
out3 = curr_dir()

```

### Task 1.1-H: Klassendefinitionen (5P)

Another important aspect which you will often find useful is the definition of classes. Instead of a function you define a structure which needs to be initialized and provides certain functionalities. The idea is that the class contains properties which can be accessed at various points inside the code/programm. This is something which a simple function cannot provide. In practice a class must contain an `__init__(self, *args)` method and can contain various other methods. Start to implement the following: 
- Create a class Animal that requires the attributes color and age (in years) during initialization.
- Write a class function `return_color(self)`, which returns the color value
- Write a class function `return_age_in_hours(self)`, which returns the age in hours (assume a year has 365 days).
- Write another class function `set_age(self, years)`, which overwrites the age value set during initialization.
- Create a class instance `mike` from the class Animal, which is "gray" and 10 years old.

<!-- grade_id: 11-h-answer, points: 0, grade: false, solution: true -->
```python
class Animal():
    # YOUR CODE HERE
    def __init__(self,color,age):
        self.color = color
        self.age = age
    def return_color(self):
        return self.color
    def return_age_in_hours(self):
        return self.age*365*24
    def set_age(self,years):
        self.age = years
    
# YOUR CODE HERE
mike = Animal("grey",10)
```

<!-- grade_id: 11-h-test, points: 5, grade: true, solution: false -->
```python
assert hasattr(mike, "color")
assert hasattr(mike, "age")

```

### Task 1.1-I: Objektorientiertes Programmieren (5P)

An interesting and often very useful property of classes is that they can inherit properties from other classes. Thereby, a hierarchical structure can be created where a parent class, e.g., Animal provides basic functionalities, while more specific properties are implemented in the child class. Let's try it out:
- Implement a class `Dog` that inherits from `Animal`. Introduce a third argument `height` (in cm) during intialization that is stored as a class variable
- Overwrite the function `return_age_in_hours` to have two return arguments: First the age in years and second the age in hours.
- Create an instance `cooper` that is yellow, 4 years old and 50 cm high.

<!-- grade_id: 11-i-answer, points: 0, grade: false, solution: true -->
```python
class Dog(Animal):
    # YOUR CODE HERE
    def __init__(self,color,age,height):
        super().__init__(color,age)
        self.height = height
    
    def return_age_in_hours(self):
        age_year = self.age
        age_hour = self.age*365*24
        return age_year, age_hour
    

# YOUR CODE HERE
cooper = Dog("yellow",4,50)
```

<!-- grade_id: 11-i-test, points: 5, grade: true, solution: false -->
```python
assert hasattr(cooper, "height")
age_y, age_h = cooper.return_age_in_hours()

```

```python

```